Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **32 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Simon Andren"
COLLABORATORS = "N/A"
```

---

# CS110 Pre-class Work - Maximum subarray problem—part II

## Question 1 [time estimate: 2 minutes]
Paste in your Python implementation of the maximum subarray from the previous class in the cell below and use that to find out the value of the maximum subarray of this array: `A = [-2, -3, 4, -1, -2, 1]`

<!-- grade_id: cell-e69c8bbcd40ca242, points: 0, grade: true, solution: true -->
```python
def bruteforce_max_subarray(A):
    """
    Implements brute-force maximum subarray finding.
    
    Parameters
    ----------
    A : list
        a NON-EMPTY list of floats
    
    Returns
    -------
    tuple
        - the start index of the max subarray
        - the end index of the max subarray
        - the value of the maximum subarray
        
    """
    
    start_ind = 0 #Define the start index for the array 
    end_ind = len(A)-1 #Define the end index for the array
    max_value = -float('inf') #Assignes the first max-value. Which will be the sum of the first
                                                #and second item in the array
    
    for i in range(start_ind, end_ind+1): #loops over all the items in the list, including the last item

        for j in range(i, end_ind+1): #Loops over all the item beyond the ith item (including the ith and last item)
            
            if sum(A[i:j+1]) > max_value: #If the sum of the subarray i to j (inclusive j) is larger than the max_arrray
                
                max_value = sum(A[i:j+1]) #Reassign the value of the max array
                start_ind = i #Reassign the start index of the max array
                end_ind = j #Reassign the end index of the max array
        
    
    #Return the start_ind, end_ind and max_value
    return (start_ind, end_ind, max_value)

    raise NotImplementedError()
    
A = [-2, -3, 4, -1, -2, 1]
bruteforce_max_subarray(A)
```

Output:
```
(2, 2, 4)
```

## Question 2 [time estimate: 5 minutes]
Now, your friend Joe comes and appends a single extra number at the end of the array, which becomes:  `B = [-2, -3, 4, -1, -2, 1, 8]`. Do you have to re-run the entire maximum subarray again? Explain your answer. 
The subsequent questions will help you figure out an efficient algorithmic strategy to address the last question, but make sure to write your explanation above first, before answering the remaining questions.


<!-- grade_id: cell-d4b7cd845c816263, points: 0, grade: true, solution: true -->
I would be faster to run the computer algorithm to find out than to find it out by ourself. However,  we would not need to re-run the entire maximum subarray again. 

Instead, we know that the maximum subarray would either exist in the first half of the array, the second half of the array or in an array that crosses the mid-point. We know that the maximum array for the first half is for, which we found out earlier (4). We can prety easy see that the maximum subarray in the second half is 9 (8+1). The largest subarray for an array that crosses the middle, would be the sum of the largest array that either end or start at the middle. For an array that end at the middle we have 4 as the maximum array, and for an array that starts at the middle, the largest array is 6. Summing together those maximum arrays, we are given a total sum of 10, which will be highest array and such the solution to the maximum array problem

## Question 3 [time estimate: 5 minutes]

**Determine if the following statement is True or False and explain your answer.**
If the maximum subarray of the array A is different than the maximum subarray of the array B (questions 1 and 2), the new maximum-subarray doesn’t need to include 8 (i.e., the newly appended element). 


<!-- grade_id: cell-9377964a8756f13b, points: 0, grade: true, solution: true -->
False! 

If the new input is a positive integer, a differing maximum subarray would include 8. This is because an 8 will always poositivly contribute to the maximum sum of an array. The maximum array in A can either be found as it is in B and in that case, they sums wouldn't be any different. But the maximum array can also be found outside that array, however, that would need to be including 8 because the "cutting" of that 8 can make in an array, will always positivly contribute to an subarray. This would not be true if the added number were a negative integer.

## Question 4 [time estimate: 10 minutes]
Complete the Python function `incremental_max_subarray(x, mx)` in the cell below.

This [video](https://www.youtube.com/watch?v=AAgErqQmwmA&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=4&t=0s)  might be helpful to understand the `incremental_max_subarray` problem.

<!-- grade_id: cell-0230e459f3d701f8, points: 0, grade: false, solution: true -->
```python
def incremental_max_subarray(x, mx):
    """ 
    Parameters
    ----------
    x : list
        A NON-EMPTY list of numbers (e.g., the 
        array B in the first two questions above)
        
        If x has 1 element: returns the value of 
        the element regardless of the value of mx
        
    mx 
        The maximum subarray of x excluding its 
        ast element (i.e., compute the maximum 
        subarray of the input array x considering 
        only its first len(x) - 1 elements)
    
    Returns
    -------
    int
        The maximum subarray of the array x
    
    Example
    -------
    Using the array B in question 2, the result 
    of incremental_max_subarray(B, 4) is 10: 
    
        10 = 8 + 1 - 2 -1 + 4 
    
    
    """
    
    if len(x) == 1: #If the array only contain a single element, return that element
        return x[0]
    
    #Initialize a current sum of 0
    x_sum = 0
    
    #Initiialize a for-loop that loops through all the possible arrays that include the last element
    for ind in range(len(x)-1,0-1,-1): #from the end index, to 0, decreasing with -1 each loop
        
        #Appends the value of current item on the current index to the running sum
        x_sum += x[ind]
        
        #It the current array (including the last element) is larger than the previous maxarray
        if x_sum >= mx:
            mx = x_sum #Change the value of the maximum sum
    
    #Return the current maximum array sum
    return mx

    raise NotImplementedError()
```

<!-- grade_id: cell-9abf19e56620ffa3, points: 1, grade: true, solution: false -->
```python
B = [-2, -3, 4, -1, -2, 1, 8]
assert(incremental_max_subarray(B, 4) == 10)
assert(incremental_max_subarray(B[:1], 0) == B[0])
```

## Question 5

### Part 1 [time estimate: 5 minutes]

Now use `incremental_max_subarray(x, mx)` iteratively, starting from the first element, to calculate the maximum subarray of a given list. Come up with a few test cases to test your code.

<!-- grade_id: cell-fd96d4ccccd99fe6, points: 0, grade: false, solution: true -->
```python
def max_subarray(A):
    """
    Using `incremental_max_subarray` iteratively on A 
    to produce the value of the maximum subarray of A.
    
    Parameters
    ----------
    A : list
        A NON-EMPTY list of floats
        
    Returns
    -------
    float
        The sum of the maximum subarray of A 
    
    """
    mx = -float('inf')
    
    #initialize a for-loop that loop over the array
    for ind in range(0, len(A)):
        
        #Assigns the returning sum of the maximum array to mx.
        mx = incremental_max_subarray(A[0:ind+1], mx) #calls on incremental max subarray with a slice of array A
        #The slice increases and increase to eventually include the entire array 
    
    #raise NotImplementedError()
    return mx

A = [-2, -3, 4, -1, -2, 1, 8]


print(max_subarray(A) == bruteforce_max_subarray(A)[2])
print(max_subarray(A[0:5]) == bruteforce_max_subarray(A[0:5])[2])
print(max_subarray(A[:1]) == A[0])
```

Output:
```
True
True
True
```

<!-- grade_id: cell-669ad779766aa2c3, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Part 2 [time estimate: 5 minutes]
Is this more efficient than the divide-and-conquer approach? Explain.

<!-- grade_id: cell-cd8f0b130a7136db, points: 0, grade: true, solution: true -->
I believe that the divide and conquer method is more efficient. As we saw in the reading, the run-time for the divide and conquer method is $\theta(n*log(n))$. For the incremental appoach, we will have, instead of log(n) layers, n layers. Each layer will i time to compute (depending on how many iterations it has to make). Because of this, I believe that the run-time for the iterative approach will be $\theta(n^2)$ which is higher than $\theta(n*log(n))$ 

```python

```